---
title: 'ทำความรู้จัก R'
description: 'Chapter description goes here.'
---

## R 101

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: aa27b0c47e
xp: 100
```

**Data Types**

สำหรับ data types พื้นฐานใน R จะมีทั้งหมด 3 แบบคือ numeric, character และ logical

numeric คือตัวเลขแบบมีทศนิยม หรือเป็นตัวเลขจำนวนเต็มก็ได้ เช่น 3.45 2.00 100 เป็นต้น

character คือตัวอักษรหรือภาษาโปรแกรมมิ่งอื่นๆเรียกว่า string เช่น "Hello World" "Awesome" "R Programming"

logical เหมือนการตอบคำถาม yes/no question ใน R เราใช้ TRUE และ FALSE เพื่อนำเสนอข้อมูลที่เป็น logical 

`@instructions`
ลองสร้างตัวแปร x, y, z เป็น numeric, character และ logical ตามลำดับ

`@hint`
ตัวอย่างเช่น

x <- 100
print(class(x))

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# numeric

# character 

# logical

```

`@solution`
```{r}

```

`@sct`
```{r}

```
