---
title: 'ทำความรู้จัก R'
description: 'Chapter description goes here.'
---

## R 101

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: aa27b0c47e
xp: 100
```

**Data Types**

สำหรับ data types พื้นฐานใน R จะมีทั้งหมด 3 แบบคือ numeric, character และ logical

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# numeric

# character 

# logical

```

`@solution`
```{r}

```

`@sct`
```{r}

```
